# Example exercise

This is an example exercise notebook.  First, a code cell:

```python
import sys
from py2neo import neo4j, node, rel

# create connection to database using defaults
graph_db = neo4j.GraphDatabaseService()

# create connection to database using explicit URL of REST API endpoint.
graph_db = neo4j.GraphDatabaseService( "http://localhost:7474/db/data/" )
```

How does it make you feel that we placed non-working and totally unrelated code at the top of this exercise?

<!-- grade_id: how_does_that_make_you_feel, points: 1, grade: true, solution: true -->


## Part A (2 points) - squares

Write a function that returns a list of numbers, such that xi=i2, for 1≤i≤n. Make sure it handles the case where n<1 by raising a ValueError.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares( n ):
    """
    Compute the squares of the numbers one to n, such that the
    ith element of the returned list equals i^2.
    """
    
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError( "n must be greater than or equal to 1" )
    return [ i ** 2 for i in range( 1, n + 1 ) ]
    ### END SOLUTION
```

Your function should print [1, 4, 9, 16, 25, 36, 49, 64, 81, 100] for n=10. Check that it does:

```python
squares( 10 )
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
from nose.tools import assert_equal
assert_equal( squares( 1 ), [ 1 ] )
assert_equal( squares( 2 ), [ 1, 4 ] )
assert_equal( squares( 10 ), [ 1, 4, 9, 16, 25, 36, 49, 64, 81, 100 ] )
assert_equal( squares( 11 ), [ 1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121 ] )
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an erorr for invalid inputs"""
from nose.tools import assert_raises
assert_raises( ValueError, squares, 0 )
assert_raises( ValueError, squares, -1 )
```

## Part B (1 point) - sum of squares

Using your squares function, write a function that computes the sum of the squares of the numbers from 1 to n. Your function should call the squares function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    ### BEGIN SOLUTION
    return sum(squares(n))
    ### END SOLUTION
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares( 10 )
```

<!-- grade_id: correct_sum_of_squares, points: 0, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert_equal(sum_of_squares(1), 1)
assert_equal(sum_of_squares(2), 5)
assert_equal(sum_of_squares(10), 385)
assert_equal(sum_of_squares(11), 506)
```

<!-- grade_id: sum_of_squares_uses_squares, points: 0, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    assert_raises(NameError, sum_of_squares, 1)
except AssertionError:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```