Import statements go at the top of the notebook, to make it clear what libraries are required:

```python
%matplotlib inline

import numpy as np
import matplotlib.pyplot as plt
```

## Part A (2 points)

There should be a header for each part of the problem. Use level two headers (##) and include the number of points this part of the problem is worth.

Give generic instructions/background in normal markdown cells.

<div class="alert alert-success">
For specific instructions, put them in a green box in a separate markdown cell. Then, have the next cell be an answer cell where students need to put their solutions:
</div>

<!-- grade_id: my_function, points: 0, grade: false, solution: true -->
```python
def my_function(x, n):
    """
    Give detailed instructions about the parameters and return values in
    the function docstring. Use the numpy docstring style to do this.
    
    Parameters
    ----------
    x : int or float
        A number
    n : int
        Another number to multiply x by
        
    Returns
    -------
    out : int
        A number equivalent to x * n

    """
    ### BEGIN SOLUTION

    # use the special syntax to delimit the region where students will
    # actually write their answer and then put in an example solution
    return x * n
    
    ### END SOLUTION
```

Sometimes, it is helpful to students to give them a few examples of how to call the function. You can also leave them  a cell instructing them to write their own test cases and explore playing around with the function.

```python
# add your own test cases here!

```

<!-- grade_id: test_my_function, points: 2, grade: true, solution: false -->
```python
"""Test that my_function works correctly."""
from numpy.testing import assert_almost_equal

# numpy.testing.allclose is also really useful (for comparing numpy arrays)

# test a few different outputs; if using floats, use an assert
# function that allows room for floating point error
assert_almost_equal(my_function(1, 1), 1)
assert_almost_equal(my_function(1.2, 1), 1.2)
assert_almost_equal(my_function(3, 4), 12)
assert_almost_equal(my_function(10, 20), 200)

# test a lot of different things so they can't just use a bunch
# of if statements
for i in np.arange(0, 10, 0.1):
    assert_equal(my_function(i, 1), i)
    assert_equal(my_function(i, 2), i + i)
    assert_equal(my_function(i, 3), i + i + i)

# print success so students know they have passed the tests
print('Success!')
```

Output:
```
Success!
```

---
## Part B (1 point)

When starting a new part of the problem, add a horizontal line before the header to visually separate the different parts.

<div class="alert alert-success">
For written answers, you should still include the instructions in a green box in its own markdown cell.
</div>

<!-- grade_id: part_b, points: 1, grade: true, solution: true -->
This is the answer to the written question. You can write the solution or rubric here.