# Ex 1. Interacting with online services

In this exercise, we will perform some basic requests obtaining information from the online services.

## Table of Contents

- [Part A: Getting online](#Part-A:-Getting-online)

- [Part B: Obtaining metadata from Crossref](#Part-B:-Obtaining-metadata-from-Crossref)

- [Part C: Obtaining a list of DOIs](#PART-C:-Obtaining-a-list-of-DOIs)

Get everything ready before we start the exercise.

```python
%load_ext autoreload
%autoreload 2
import sys
sys.path.append("../modules/orcid-python")
sys.path.append("../modules/pyalm")

import os
#os.environ['PLOS_API_KEY'] = 'user api key'

import requests
import time
import orcid

# import json for output
import json
```

## Part A: Getting online

Back to [Table of Contents](#Table-of-Contents).

The first part of the exercise is to simply to successfully request a page online and recieve back a 200 HTTP response. We imported the requests library above and will download the google.com page. This exercise is largely to ensure that you have a properly functioning internet connection! 

<div class="alert alert-success">
An example of using the `requests` library is shown in the first example notebook from the book chapter ("1. Programming against an API - Crossref"). Adapt that code to download the homepage of google.com, store the result of getting the google.com home page in a variable named `response`, and show that the HTTP status response code is 200.  
</div>

<!-- grade_id: A-Get_URL, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# First give the URL that you want to obtain
url = 'http://google.com'

# Then set up the request. Look at example Notebook #1 for this session and copy the names of object that the
# request response will go into.
response = requests.get(url)

# output the response code
print( "response status code: " + str( response.status_code ) )
### END SOLUTION
```

<!-- grade_id: A-Get_URL-test, points: 1, grade: true, solution: false -->
```python
# test
assert response.status_code == 200 # If you get the code above working this assertion should pass
```

## Part B: Obtaining metadata from Crossref

Back to [Table of Contents](#Table-of-Contents).

Now we have a working connection we will obtain some bibliographic metadata from Crossref. In the example notebook we obtained data from a single DOI. First we will replicate that for a different DOI 

<div class="alert alert-success">
Obtain the Crossref metadata for the DOI 10.1038/171740a0 Obtain the names of the authors of this article and its title.  You'll be using `requests.get()` to make your request.  Make sure to store the result of that call in a variable named `response`.
</div>

<!-- grade_id: B-Get_metadata, points: 0, grade: false, solution: true -->
```python
# Now write the code the obtain the relevant JSON metadata, check that the status code is 200.

### BEGIN SOLUTION
# In a second cell below uncomment the code to look at the JSON output, then select the author's surnames to print

# start with the domain name of the server
BASE_URL = "http://api.crossref.org"

# add on details of exact aPI we are referencing.
API_action = "works"
doi_to_lookup = "10.1038/171740a0"

# tie together into URL.
url = BASE_URL + "/" + API_action + "/" + doi_to_lookup

print( "My URL = " + url )

# access URL.
response = requests.get( url )
### END SOLUTION
```

<!-- grade_id: B-Get_metadata-test, points: 1, grade: true, solution: false -->
```python
# test - status code should be 200
assert response.status_code == 200
```

```python
# Convert the response to a JSON object and see what we got!
root_json_object = response.json()

# output 
print json.dumps( root_json_object, sort_keys = True, indent = 4, separators = ( ',', ': ' ) )
```

<div class="alert alert-success">
When you call `response.json()`, the structured JSON data in the response is converted into a combination of dictionaries (each JSON object, composed of name-value pairs, becomes a python dictionary) and lists (each JSON list of values becomes a python list).  Use your knowledge of python dictionaries and lists to retrieve the surnames of the two authors from the data structure you printed above and store the first in variable `surname1` and the second in `surname2`, and retrieve the title of the article and store it in the variable `title`.  Then print all three variables out.
</div>

<!-- grade_id: B-Get_article_info, points: 0, grade: false, solution: true -->
```python
# Now use your knowledge of python dictionaries to print the surname of each of the authors and its title

### BEGIN SOLUTION

# find author information - all in one line.
#surname1 = root_json_object['message']['author'][0]['family']
#surname2 = root_json_object['message']['author'][1]['family']
#title = root_json_object['message']['title'][0]

# NOTE FOR EHMAD/CHRISTINA - The authors surnames are Franklin and Gosling (possibly in capitals).
# This is the Rosalind Franklin article that provided the evidence backing up Watson and Crick's
# paper on the structure of DNA. Not sure how to set up the marking for this but they should be
# able to pull out the surnames by inspecting the JSON.

# start with the domain name of the server
#BASE_URL = "http://api.crossref.org"

# add on details of exact aPI we are referencing.
#API_action = "works"
#doi_to_lookup = "10.1038/171740a0"

#url = BASE_URL + "/" + API_action + "/" + doi_to_lookup

#print( "My URL = " + url )

#response = requests.get( url )
#root_json_object = response.json()

#print( "status, first level = " + root_json_object[ "status" ] )

# first get message object
message_dict = root_json_object[ "message" ]

#print( "message_dict:" )
#print json.dumps( message_dict, sort_keys = True, indent = 4, separators = ( ',', ': ' ) )

# get author list
author_list = message_dict[ "author" ]

#print( "author_list:" )
#print json.dumps( author_list, sort_keys = True, indent = 4, separators = ( ',', ': ' ) )

# get author 1
author_1_dict = author_list[ 0 ]

print("\n\nauthor 1:" )
print json.dumps( author_1_dict, sort_keys = True, indent = 4, separators = ( ',', ': ' ) )

surname1 = author_1_dict[ "family" ]

# get author 2
author_2_dict = author_list[ 1 ]

print("\n\nauthor 2:" )
print json.dumps( author_2_dict, sort_keys = True, indent = 4, separators = ( ',', ': ' ) )

surname2 = author_2_dict[ "family" ]

# retrieve title
# note - "title" contains a list, with one item in it.
#     You can't just get  "title":
#title = message_dict[ "title" ]

# you need to get the title list, then get the first item in it.
title_list = message_dict[ "title" ]
title = title_list[ 0 ]

# Once you've retrieved the values, print them out!
print( "\n\n" )
print( "Surname 1 = " + surname1 )
print( "Surname 2 = " + surname2 )
print( "Article Title = " + title )

# rather than reference each author using indices, you could
#    also loop over the author list.
print( "\n\nLoop over authors:" )
author_counter = 0
for author in author_list:
    
    # increment author counter
    author_counter += 1
    
    # get surname out of current author's dictionary.
    current_surname = author[ "family" ]

    # print it.
    print( "- Author #" + str( author_counter ) + " ( index " + str( author_counter - 1 ) + " ) surname = " + current_surname )

#-- END loop over authors --#

print( "\n\nentire JSON object:" )
print json.dumps( root_json_object, sort_keys = True, indent = 4, separators = ( ',', ': ' ) )

### END SOLUTION

# Once you've retrieved the values, print them out!
print( "Surname 1 = " + surname1 )
print( "Surname 2 = " + surname2 )
print( "Article Title = " + title )
```

<!-- grade_id: B-Get_article_info-test, points: 1, grade: true, solution: false -->
```python
# test that names are what we expect
assert surname1.lower() == "franklin"
assert surname2.lower() == "gosling"
assert title.lower() == "molecular configuration in sodium thymonucleate"
```

## PART C: Obtaining a list of DOIs

Back to [Table of Contents](#Table-of-Contents).

In this part of the exercise you will do something not covered in the class notes, obtain a set of DOIs from the Crossref search API. In the notes we show collecting DOIs from ORCID. In the second set of exercises we will show an example using a publisher API. 

To prepare for this you will need to look at the API documentation at https://github.com/CrossRef/rest-api-doc/blob/master/rest_api.md. In particular look at the section on queries. You will see that quite sophisticated queries are possible. You may want to test some examples and look at the results below. Think carefully about how you want to construct the URL for your query as you do this.

<div class="alert alert-success">
Use the CrossRef search API to search for articles by Martin Karplus.  Make a list of the DOIs of each of the results, stored in a variable named `dois`.  Print out the contents of this list.  You should have 20 matches, and so 20 DOIs (the API's default limit of matches returned for a query).
</div>

<!-- grade_id: C-crossref_search-request, points: 1, grade: false, solution: true -->
```python
# Construct a URL for your search, submit it to the API, then call repsonse.json()
#    to convert the JSON response to python objects and store the result in a
#    variable named json_response.

### BEGIN SOLUTION
url = 'http://api.crossref.org/works?query=Martin+Karplus'
response = requests.get(url)
json_response = response.json()
### END SOLUTION
```

<!-- grade_id: C-crossref_search-doi_list, points: 0, grade: false, solution: true -->
```python
# Now collect the DOIs from the JSON response into a list stored in variable "dois",
#    and print the DOIs.
dois = []

### BEGIN SOLUTION
for item in json_response['message']['items']:

    current_doi = item[ 'DOI' ]
    dois.append( current_doi )
    print( "Current DOI: " + str( current_doi ) )

### END SOLUTION
    
print( "Found " + str( len( dois ) ) + " DOIs!" )
```

<!-- grade_id: C-crossref_search-doi_list-test, points: 1, grade: true, solution: false -->
```python
# test - There should be 20 DOIs because that is the default number of results that the Crossref API returns
assert len(dois) == 20
```